# Game theory in Python

In this notebook we reiterate the topics touched in the lecture, get used to the basic concepts of game theory, as well as how to implement it using Python. The notebook contains four parts:
- The (iterated) prisonners dilemma,
- Evolutionary hawk-dove game,
- Calculating Nash equilibria,
- Calculating dominant strategies.

First we install the package axelrod and test if we can import it.

```python
from IPython.display import clear_output
!pip install axelrod 
import axelrod
clear_output()
print("Everything A-Okay!")
```

### Effective Choice in the Prisoners Dilemma
In the first part of the notebook we will recreate the research originally done by Axelrod. Start with reading his [paper](AxelrodEffectiveChoiceinthePrisonersDilemma.pdf).

Fortunately for us, the prisoners dilemma is already implemented as a Python package, called Axelrod. Take a look at their [GitHub page](https://github.com/Axelrod-Python/Axelrod) and their [readthedocs](https://axelrod.readthedocs.io/en/stable/).

<table>
    <tr>
        <td></td>
        <td></td>
        <td>Player A</td>
        <td></td>
    </tr>
    <tr>
        <td></td>
        <td></td>
        <td>C</td>
        <td>D</td>
    </tr>
    <tr>
        <td>Player B</td>
        <td>C</td>
        <td>(3, 3)</td>
        <td>(0, 5)</td>
    </tr>
    <tr>
        <td></td>
        <td>D</td>
        <td>(5, 0)</td>
        <td>(1, 1)</td>
    </tr>
</table>


To get used to the workflow of axelrod we begin by implementing a simple match. In this match we will want to see the result of a prisonners dilemma between a so-called cooperator strategy, and a so-called alternator strategy. We already implemented this for you:

```python
import axelrod as axl

# make a tuple, called players, containing Cooperator and an Alternator
players = (axl.Cooperator(), axl.Alternator())

# let them play
match = axl.Match(players, turns=5)
turns = match.play()

# extract the final_score and winner from the match
final_score = match.final_score()
winner = match.winner()

print(f'turns: \t\t{turns}')
print(f'final_score: \t{final_score}')
print(f'winner: \t{winner}')
```

Axelrod 'organised' a tournament with many players. Here we will recreate the (first) tournament with axelrod. Take a look at [which strategies of the (first) tournament are implemented in axelrod](https://axelrod.readthedocs.io/en/stable/reference/overview_of_strategies.html#axelrod-s-first-tournament), and take a look [how to implement a tournament](https://axelrod.readthedocs.io/en/stable/tutorials/getting_started/tournament.html).


```python
strategies = [axl.Random(), 
              axl.TitForTat(),
              axl.FirstByTidemanAndChieruzzi(),
              axl.FirstByNydegger(),
              axl.FirstByGrofman(),
              axl.FirstByShubik(),
              axl.FirstBySteinAndRapoport(),
              axl.Grudger(),                        # Friedman
              axl.FirstByDavis(),
              axl.FirstByGraaskamp(),  # different implementation than in original paper
              axl.RevisedDowning(),    # bug in original implementation, when revised best strategy!
              axl.FirstByFeld(),
              axl.FirstByJoss(),
              axl.FirstByTullock(),
              axl.FirstByAnonymous()]

tournament = axl.Tournament(strategies, 
                            turns=200,
                            repetitions=10)


results = tournament.play()
plot = axl.Plot(results)

plot.boxplot();
plot.payoff();
```

Notice the similarity in strategy between the high scoring strategies? They are never the first to defect, **be nice**!

### Evolutionary hawk-dove game
The hawk-dove game is a game where two birds fight over some food, where the reward of the food is V. If both birds decide to be unagressive (dove), they each get half of the food. If one of the birds decides for fight for the food (hawk), while the other bird is unaggresive, the hawk gets all the food. If both birds are aggresive, they each get half the food minus the cost (C) of fighting:

<table>
    <tr>
        <td></td>
        <td></td>
        <td>Player A</td>
        <td></td>
    </tr>
    <tr>
        <td></td>
        <td></td>
        <td>Dove</td>
        <td>Hawk</td>
    </tr>
    <tr>
        <td>Player B</td>
        <td>Dove</td>
        <td>(V/2, V/2)</td>
        <td>(0, V)</td>
    </tr>
    <tr>
        <td></td>
        <td>Hawk</td>
        <td>(V, 0)</td>
        <td>((V-C) / 2, (V-C) / 2)</td>
    </tr>
</table>

In this assignment we will implement an evolutionary hawk-dove game, as well as calculate the expected payoff for each strategy. 

$$
\begin{cases}
(V - C) / 2 > 0 & \text{Hawk strategy dominant}\\
(V - C) / 2 < 0 & \text{Mixed strategy!}
\end{cases}
$$

For a mixed strategy to hold the expected payoff for Hawk must be equal to Dove:
$$
E(U_D) = E(U_H)
$$

The expected payoff for dove:
$$
E(U_D) = \frac{V}{2}(1 - p_H)
$$

The expected payoff for hawk:
$$
E(U_H) = \frac{V - C}{2}p_H + V (1-p_H)
$$

Now solve:
$$
\frac{V}{2}(1 - p_H) = \frac{V - C}{2}p_H + V (1-p_H) \\ 
$$
And implement function proportion_hawk, that calculates the proportion of hawks according to this equation.

<!-- grade_id: cell-e84fed0f9475b468, points: 0, grade: false, solution: true -->
```python
def proportion_hawk(V, C):
    """
    This function returns the proportion of hawks for given V and C.
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-75990e61db18b727, points: 1, grade: true, solution: false -->
```python
assert isinstance(proportion_hawk(1, 2), (int, float)), "The function should return numerical (int of float) values"
assert 0 <= proportion_hawk(1, 2) <= 1, "A proportion value is always between 0 and 1!"
assert proportion_hawk(2, 1) == 1, "This is a hawk-dominant solution!"
```

We can test if our calculation is correct/holds in an evolutionary setting. We will test this by implementing a [Moran Process](https://axelrod.readthedocs.io/en/stable/tutorials/getting_started/moran.html), which is an evolutionary game, where after each round in the tournament one random individual changes its strategy based on fitness proportionate selection. Implement your a [new game](https://axelrod.readthedocs.io/en/stable/tutorials/advanced/games.html) based on the hawk-dove payoff, and feed it to a Moran process.

<!-- grade_id: cell-963495b08cefad96, points: 0, grade: false, solution: true -->
```python
import matplotlib.pyplot as plt

# Set the values 
V = 1
C = 3

# Generate for instance 25 hawks (defectors) and 25 (cooperators)
# YOUR CODE HERE
raise NotImplementedError()

# Implement a hawk dove game
# YOUR CODE HERE
raise NotImplementedError()

# feed the hawkdove game to the MoranProcess
mp = axl.MoranProcess(players, turns=1, game=hawkdove)

# lets see how it evolves
iters = 1000
for _ in range(iters):
    if not mp.fixation_check():
        next(mp)
        
# now make a pretty plot
mp.populations_plot()
plt.show()
```

<!-- grade_id: cell-8378b6b057fb7c64, points: 1, grade: true, solution: false -->
```python
assert "players" in vars(), "store the the strategies as variable players"
assert all(type(player).__name__ in ['Cooperator', 'Defector'] for player in players), "only hawks and doves, please!"

assert "hawkdove" in vars(), "store the hawk dove ´game' as variable hawkdove"
assert type(hawkdove).__name__ == 'Game'

```

### Nash equilibria
Nash equilibria are equilibria where no player has anything to **gain** by changing only their own strategy. Implement the function `find_nash()`, that returns the list of nash equilibria for a payoff matrix. We already implemented the functions max_row and max_col for you, which might help in solving the assignment. 

<!-- grade_id: cell-e7b353ff5650f7f9, points: 0, grade: false, solution: true -->
```python
def max_row(matrix, row_nr):
    """
    A helper function that returns the maximum payoff for player B by column index
    """
    return max(x[0] for x in matrix[row_nr])

def max_col(matrix, col_nr):
    """
    A helper function that returns the maximum payoff for player A by row index
    """
    return max(matrix[i][col_nr][1] for i in range(len(matrix)))
    
def find_nash(payoff_matrix):
    """
    Returns a list of indices of the Nash-equilibria. 
    If no indices are present returns an empty list.
    """
    nash_equilibria = []
    # YOUR CODE HERE
    raise NotImplementedError()
    return nash_equilibria
```

<!-- grade_id: cell-1c0c928bc993bd0f, points: 2, grade: true, solution: false -->
```python
payoff_matrix = [[(2, 2), (3, 1)],
                 [(1, 1), (0, 0)]]
assert set(find_nash(payoff_matrix)) == set([(0, 1)])

payoff_matrix = [[(1, 0), (0, 1)],
                 [(0, 1), (1, 0)]]
assert set(find_nash(payoff_matrix)) == set([])

payoff_matrix = [[(2, 2), (0, 0)],
                 [(0, 0), (1, 1)]]
assert set(find_nash(payoff_matrix)) == set([(0, 0), (1, 1)])

payoff_matrix = [[(0, 0), (1, 1)],
                 [(1, 1), (1, 1)]]
assert set(find_nash(payoff_matrix)) == set([(0, 1), (1, 0), (1, 1)])

payoff_matrix = [[(0, 0), (1, 1), (2, 2)],
                 [(1, 1), (1, 1), (1, 1)]] 
assert set(find_nash(payoff_matrix)) == set([(1, 0), (1, 1), (0, 2)])

```

### Iterated Elimination of Dominated Strategies (IEDS)
A strategic dominant strategy is a strategy that is always better than another strategy for a player. Lets assume the following payoff table: 
<table>
    <tr>
        <td></td>
        <td></td>
        <td>Player A</td>
        <td></td>
    </tr>
    <tr>
        <td></td>
        <td></td>
        <td>C</td>
        <td>D</td>
    </tr>
    <tr>
        <td>Player B</td>
        <td>C</td>
        <td>(1, 2)</td>
        <td>(2, 1)</td>
    </tr>
    <tr>
        <td></td>
        <td>**D**</td>
        <td>**(4, 1)**</td>
        <td>**(3, 2)**</td>
    </tr>
</table>
For player A it is unclear whether to play Cooperate or Defect, since it depends on the strategy of the other player which is best. However, for player B it is always better to defect, since those rewards are always higher than the cooperation scores. We say that Cooperation is a (strictly) dominated rule, we can therefore ignore this row in the game (under the assumption that the player is purely rational):
<table>
    <tr>
        <td></td>
        <td></td>
        <td>Player A</td>
        <td></td>
    </tr>
    <tr>
        <td></td>
        <td></td>
        <td>C</td>
        <td>**D**</td>
    </tr>
    <tr>
        <td>Player B</td>
        <td></td>
        <td></td>
        <td></td>
    </tr>
    <tr>
        <td></td>
        <td>D</td>
        <td>(4, 1)</td>
        <td>**(3, 2)**</td>
    </tr>
</table>
Since player A knows the payoff for player B, it knows that player B will always play Defect. Using this knowledge, it is obvious for player A that Defection always has a higher reward than Cooperation (Cooperation is dominated by Defect). We can therefore also ignore the Cooperation column for player A. Now we end up with a simple table with one **dominant** option:
<table>
    <tr>
        <td></td>
        <td></td>
        <td>Player A</td>
        <td></td>
    </tr>
    <tr>
        <td></td>
        <td></td>
        <td></td>
        <td>D</td>
    </tr>
    <tr>
        <td>Player B</td>
        <td></td>
        <td></td>
        <td></td>
    </tr>
    <tr>
        <td></td>
        <td>D</td>
        <td></td>
        <td>(3, 2)</td>
    </tr>
</table>
Rational players will always always play defect with the original payoff matrix.

In the next assignment we will implement the Iterated Elimination of Dominated Strategies algorithm. The algorithm works by iteratively removing dominated strategies, just like we did in the example. You only have to remove **strictly dominated** strategies! We already implemented the functions `remove_row()` and `remove_col()` for you, which might help in solving the assignment.

<!-- grade_id: cell-12ce60a0d624b2ed, points: 0, grade: false, solution: true -->
```python
import copy

def remove_row(matrix, row_nr):
    """
    A helper function that deletes a row by its index
    """
    del matrix[row_nr]
    
def remove_col(matrix, col_nr):
    """
    A helper function that deletes a column by its index
    """
    for row_i in range(len(matrix)):
        del matrix[row_i][col_nr]


def IEDS(payoff_matrix):
    """
    The Iterated Elimination of Dominated Strategies algorithm. 
    Returns the smallest subset of non-dominated strategies
    """
    matrix = copy.deepcopy(payoff_matrix)
    # YOUR CODE HERE
    raise NotImplementedError()
    return matrix    
```

<!-- grade_id: cell-4943b85a4d67c974, points: 2, grade: true, solution: false -->
```python
payoff_matrix = [[(1, 1)]]
assert IEDS(payoff_matrix) == [[(1, 1)]]

payoff_matrix = [[(2, 2)],
                 [(1, 1)],
                 [(7, 1)],
                 [(1, 2)]]
assert IEDS(payoff_matrix) == [[(7, 1)]]

payoff_matrix = [[(2, 2), (1, 1)],
                 [(1, 1), (0, 0)]]
assert IEDS(payoff_matrix) == [[(2, 2)]]

payoff_matrix = [[(3, 1), (0, 1), (0, 0)],
                 [(1, 1), (1, 1), (5, 0)],
                 [(0, 1), (4, 1), (0, 0)]]
assert IEDS(payoff_matrix) == [[(3, 1), (0, 1)], 
                               [(1, 1), (1, 1)], 
                               [(0, 1), (4, 1)]]

```

### Spatial prisonners dilemma
The module axelrod has a built-in option for spatial tournaments, see [here](https://axelrod.readthedocs.io/en/stable/tutorials/further_topics/spatial_tournaments.html). Or you could easily combine it with mesa! Take a look at the folder mesa_axelrod for a very simple example, where there are agents walking around randomly, and playing iterated prisonners dilemma. Winners reproduce, and losers die. 